Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nazar Yaremko"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random
def worst_sort(A):
    #creating a new list that is identical to the list we want to sort
    B = A[:]
    #sorting the new list 
    B.sort()
    #randomly suffling the original list until it becomes a sorted list
    while A != B:
        random.shuffle(A)
    return A
worst_sort([2,5,1,7,9,-3,5])
```

Output:
```
[-3, 1, 2, 5, 5, 7, 9]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
In my case - big omega (1). This would happen if the algoithm luckily shuffled the list to become sorted at the first attempt. However, if we did not use sort() function, then we would have to use another sorting algorithm that can sort B and then the time complexity of the entire algorithm would depend on the sorting algorithm we choose. 

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
Big theta (N!). This is a chance that the list would be sorted, so I think that on average it would be sorted 1 out of N! times. Because we cae about the growth, the constants don't matter and the complexity would only include N!.

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
For small sizes, definitely. As the size of the list growth, the time comlexity would become massive because factorials grow very fast. 

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.

Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-tilde{δ}/2)% and (50+\tilde{δ}/2)% percentiles, where \tilde{δ}=100*δ. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


```python
def merge(A):
    #defining the base case to be 1
    if len(A)>1:
        #calculating the index of the marker that will break the list down into two sublists
        q = len(A)//2
        #creating two sublists: L and R
        L = A[:q]
        R = A[q:] 
        #recursively calling the merge function to break down and collect back both L and R lists. 
        merge(L)
        merge(R)
        
        #initializing the indexing variables
        i = j = k = 0
        #lines 18-35 comapre the elements of the created sublists like in three_way_merge1, except this time
        #we use the while loop
        while i < len(L) and j < len(R): 
            if L[i] < R[j]:
                A[k] = L[i]
                i += 1
            else:
                A[k] = R[j]
                j += 1
            k += 1
            
        while i < len(L):
            A[k] = L[i]
            i += 1
            k += 1
            
        while j < len(R):
            A[k] = R[j]
            j += 1
            k += 1
            
    return(A)
```

<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
import math
def check_approx_med(A, num, delta):
    #soring the algorithm to find the median
    merge(A)
    #finding the median
    marker = len(A)//2
    #checking of the number is withing the bounds specified by the value of delta
    #I also had to include len(A) > 1 because It seems like in any case when the length is 1, the only number in it would be the median
    #And no matter what the value of delta is, there will be no numbers left to compare
    if len(A) > 1 and num >= A[(marker - math.ceil(len(A)*delta/2))] and num <= A[(marker + math.floor(len(A)*delta/2))]:
        return True
    else:
        return False
    
check_approx_med([1,2,3,4,5,6,7,8], 5, 0.3)
```

Output:
```
True
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    #starting the count
    count = 0
    #while the count is 0, randomly choose numbers
    while count < 100//delta:
        a = random.choice(A)
        #checking if the random number is approximately a median
        if check_approx_med(A, a, delta) == False:
            count += 1
        else:
            return a, count

approx_med_finder([1,1,1,1,1,1,1,1], 0.3)
```

Output:
```
(1, 0)
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
The probability of failure would be (N- delta x N)/N which is 1 - delta. This would represent the probability of a random function choosing a number that is not within the boundaries of delta. It certainly scales with the values of delta. For instance, the smaller delta is, the bigger probability of the failurer would be. Regarding N, I don't think that N matters and no matter what N we choose (N- delta x N)/N would be the same.

x - multiplication sign

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
check_approx_med would uses the merge function to sort the list, which is nlogn. Further, it does a couple of comparisions which are constant, so for the growth of time complexity they do not matter. Further, approx_med_finder function has a probability of (delta * N)/N which is delta to find the correct number. Therefore, the overal time complexity would be nlogn + (delta x N)/N. Therefore, the overall complexity would be big theta (nlogn) which is the part of the overall complexity that grows the fastest.