Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nazar Yaremko"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    maximal_sum = float('-inf')
    # I set this to negative infinity because it is smaller than any 
    #maximum sum that could be in any list, so we could initiate the comparisons correctly
    for k in range(0,len(A)):
        #this will be out lower bound
        sum_of = 0
        #creating a veriable that will collect our sums
        for l in range(k, len(A)):
            #this will be the upper bound
            sum_of += A[l]
            #adding the numbers between k and len(A) to calculate a sum
            if sum_of > maximal_sum:
                #comparing the sum to the previous biggest sum
                maximal_sum = sum_of
                #setting the new sum as the biggest sum if the previous equation is satisfied
                low = k
                high = l
    return(low, high, maximal_sum)


bruteforce_max_subarray([-2, -3, 4, -1, -2, 1])

```

Output:
```
(2, 2, 4)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
I think that the more efficient way would be to just compare the sums between each number and the appended number, just so we needed to perform n number of steps instead of comparing all of the numbers again and again. 

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
I think that this is false. If we have the last maximum sum saved in the "maximal_sum" variable, no matter how many times we run through list A, we will have the same maximum subarray. Only by including the new number we can obtain new higher subarrays to update "maximal_sum". 

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    #checking if the len of x is 1 to satisfy the condition provided in the prompt
    if len(x) == 1:
        return x[0]
    else: 
        #establishing the intial variables
        current_max = mx
        sum_of = 0
        #creating a for loop with a negative step to compare sums of subarrays end with the last number in the list and start with numbers to the left 
        #every iteration this number shifts to the left
        for i in range(len(x)-1,-1,-1):
            sum_of += x[i]
            #comparing the current sum and substituting it with the max sum if the former one is bigger
            if sum_of > current_max:
                current_max = sum_of
        return current_max
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    maxes = []
    #iterating throught the length of the list and performing the previous function on it multiple times
    #the point is that at every iteration we capture 1 additional number untl we reach the end of the list
    for i in range(len(A)+1):
        maxes.append(incremental_max_subarray(A[0:i], A[0]))
    return max(maxes)

max_subarray([-2, -5, 6, -2, -3, 1, 5, -6])

```

Output:
```
7
```

```python
#FOR TEST
def incremental_max_subarray1(x, mx):
    count = 0
    if len(x) == 1:
        return x[0]
        count += 1
    else: 
        current_max = mx
        count += 1
        sum_of = 0
        count += 1
        for i in range(len(x)-1,-1,-1):
            sum_of += x[i]
            count += 1
            if sum_of > current_max:
                current_max = sum_of
                count += 1
        return count
```

```python
#FOR TEST
def max_subarray1(A):
    count = 0
    for i in range(len(A)+1):
        count += incremental_max_subarray1(A[0:i], A[0])
    return count

max_subarray1([-2, -5, 6, -2, -3, 1, 5, -6, 8, 1, 4, 5, 7, 2, 1, 3, 4, 5, 6, 7, 3, 2, 4, 6, 7, -5, 3, 6, 7]) 
```

Output:
```
749
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
In my opinion, no. I tried to test it with different inputs and counted steps and every time the number of steps was around n^2. Therefore, its effectivity is close to brute force. I might be completely off, but I based this answer on my observations and it seems plausible. 