Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nazar Yaremko"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 13.1

## Question 1.

Complete the following functions, following the algorithms in Cormen et al.

<!-- grade_id: cell-542b2dce56293aef, points: 0, grade: false, solution: true -->
```python
def lcs_length(x, y):
    """
    Computes the length of an LCS of strings x and y.
    
    Inputs:
    - x, y: strings
    
    Outputs:
    - c: a list of lists of ints OR a numpy array. c[i,j] contains the 
    length of a LCS of x[:i] and y[:j]
    - b: a list of lists of strings OR a numpy array, containing the information
    used for LCS reconstruction (See Cormen et al.) Use "N" (North), "NW" 
    (North West), and "W" (West) that correspond to the directions of the arrows 
    used in Cormen et al.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def lcs_length(x, y):
    #assigning sizes of our tables
    m = len(x)
    n = len(y)
    #creating the tables
    c = np.zeros([m+1, n+1])
    b = np.empty([m, n], dtype = '<S1')
    #starting the interations
    for i in range(1, m+1):
        for j in range(1, n+1):
            #if two items/chracters/numbers from the corresponding places in the strings are hte same
            #then return the combination of two appropriate subproblems
            if x[i-1] == y[j-1]:
                c[i, j] = c[i-1, j-1] + 1
                b[i-1, j-1] = "diagonal"
            #further, if they are not the same, moving in the approproate direction
            elif c[i-1, j] >= c[i, j-1]:
                c[i,j] = c[i-1,j]
                b[i-1, j-1] = "up"
            else:
                c[i,j] = c[i, j-1]
                b[i-1,j-1] = "left"
    return b, c

lcs_length('ansatz', 'algorithms')
```

Output:
```
(array([[b'd', b'l', b'l', b'l', b'l', b'l', b'l', b'l', b'l', b'l'],
        [b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u'],
        [b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'd'],
        [b'd', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u'],
        [b'u', b'u', b'u', b'u', b'u', b'u', b'd', b'l', b'l', b'u'],
        [b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u', b'u']],
       dtype='|S1'), array([[0., 0., 0., 0., 0., 0., 0., 0., 0., 0., 0.],
        [0., 1., 1., 1., 1., 1., 1., 1., 1., 1., 1.],
        [0., 1., 1., 1., 1., 1., 1., 1., 1., 1., 1.],
        [0., 1., 1., 1., 1., 1., 1., 1., 1., 1., 2.],
        [0., 1., 1., 1., 1., 1., 1., 1., 1., 1., 2.],
        [0., 1., 1., 1., 1., 1., 1., 2., 2., 2., 2.],
        [0., 1., 1., 1., 1., 1., 1., 2., 2., 2., 2.]]))
```

<!-- grade_id: cell-6b825bf900c2cf0b, points: 0, grade: false, solution: true -->
```python
def print_lcs1(b,x,i,j, lst):
    if i == 0 or j == 0:
        return None
    #here we moving from the right lower corner of table b and if the cells is b'd'
    #we add this value to the list
    if b[i-1, j-1] == b'd':
        print_lcs1(b, x, i-1, j-1,lst)
        lst.append(x[i-1])
    #otherwise we follow thee direction (either leeft or up)
    elif b[i-1,j-1] == b'u':
        print_lcs1(b, x, i-1, j, lst)
    else:
        print_lcs1(b, x, i, j-1, lst)
    return(lst)

def print_lcs(b,x,i,j):
    """
    Finds a LCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    lst = print_lcs1(b,x,i,j, [])
    return lst

print_lcs(lcs_length('algorithms', 'ansatz')[0], 'algorithms', len('algorithms'), len('ansatz'))


```

Output:
```
['a', 't']
```

<!-- grade_id: cell-e06e8f09bb8a5063, points: 1, grade: true, solution: false -->
```python
import numpy as np
x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0] == ['a', 'b', 'c'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1] == 3)

x, y = 'xyqwsssazdesaqqf', 'xoppoypllzookjdef'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0]  == ['x', 'y', 'z', 'd', 'e', 'f'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1]  == 6)

```

Output:
```
IndexError                                Traceback (most recent call last)
<ipython-input-70-7db5275a991e> in <module>()
      2 x, y = 'ambgdec', 'aubyci'
      3 c, b = lcs_length(x, y)
----> 4 assert(print_lcs(b,x,len(x)-1,len(y)-1)[0] == ['a', 'b', 'c'])
      5 assert(print_lcs(b,x,len(x)-1,len(y)-1)[1] == 3)
      6 

IndexError: list index out of range
```

## Question 2. (Adapted from Exercise 15-4.1 Cormen et al.)
Use the functions built in Question 1 to find the LCS of `'10010101'` and `'010110110'`. You should store the list that represents the LCS you found in a variable named `lcs_q2`

<!-- grade_id: cell-5d742a82dde1229f, points: 0, grade: true, solution: true -->
```python
lcs_q2 = print_lcs(lcs_length('10010101', '010110110')[0], '10010101', len('10010101'), len('010110110'))
print(lcs_q2)
```

Output:
```
['1', '0', '0', '1', '1', '0']
```

## Question 3. (Adapted from Exercise 15-4.5 Cormen et al.)
Complete the following function, making use of `lcs_length` and `print_lcs`. 



<!-- grade_id: cell-429c944f34a54c91, points: 0, grade: true, solution: true -->
```python
def lmis(lst):
    """
    Finds the Longest Monotonically Increasing Subsequence (LMIS) of a list 
    (lst) of n numbers in O(n^2) time. Note that a monotonically increasing 
    sequence is a sequence of numbers such that: a_1 <= a_2 <= ... <= a_n .
    
    Inputs:
    - lst: a list of ints
    
    Outputs:
    - out_lst: a list of ints, a longest monotonically increasing subsequence
    of lst
    """
    #here we sort the list and create a table using the original list and itself in a sorted form
    x = lst
    y = sorted(lst)
    #using the lcs_length function
    b, c = lcs_length(x, y)
    #returning the longest common subsequence which is the LMIS in this case
    return(print_lcs(b,lst,len(lst),len(lst)))
    
lmis([-1, 3, 4, 5, 2, 2, 2, 2])
```

Output:
```
[-1, 2, 2, 2, 2]
```

```python
#here i was just playing with the approach I found in once YouTube video 
def lmis2(lst):
    #here I created a small table to keep track of the subproblems
    storage_lst = [1 for i in range(len(x)+1)]
    #initializing interations
    for j in range(1, len(storage_lst)):
        for i in range(0, j):
            #here we make the appropriat comparison and update the cells 
            #these cells represent solutions to the subproblems
            if lst[j] >= lst[i]:
                storage_lst[j] = storage_lst[i]+1
            else: 
                continue
    #in the end we can see the list with the solutions
    #the largest value in the list would be the length of the LMIS
    return storage_lst

lmis2([-1, 3, 4, 5, 2, 2, 2, 2])
```

Output:
```
[1, 2, 3, 4, 2, 3, 4, 5]
```

<!-- grade_id: cell-f8b218898fb43d19, points: 0, grade: true, solution: false -->
```python
assert(lmis([5, 1, 0, 4, 2, 6, 7, 9]) == [1, 4, 6, 7, 9])
assert(lmis([6, -1,4,5,5,7,-6,2]) == [-1, 4, 5, 5, 7])
```